---
title: "Un curso amigable sobre R"
output: 
  learnr::tutorial:
    css: css/Estilo.css
runtime: shiny_prerendered
---

```{r setup, include=FALSE}

library(learnr)
library(dplyr)
library(ggplot2)
library(jpeg)
library(gridGraphics)
library(curl)
library(fontawesome)
library(showtext)
library(gradethis)
font_add_google('Gochi Hand', 'gochi')
tutorial_options(exercise.timelimit = 60, exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
```

## ¡Bienvenido!

```{r, echo = FALSE, eval = TRUE, out.width = "44%", fig.align = "center"}

knitr::include_graphics("images/Mundo_R.jpg")  
```

<section style="text-align: center;"><small>Ilustracíon hecha por <font color = "black"><u>@allison_horst</u></font>.</small></section>

### Descripción

Este curso pretende ser una introducción básica del `R`. Su objetivo consiste en enseñar habilidades en la limpieza, manipulación y visualización de datos mediante el uso del `R`. Tenga en cuenta que este curso no pretende ser un sustituto de un curso completo de programación o de estadística.

```{r eval = TRUE, echo = FALSE, fig.align = 'center', message = FALSE, fig.showtext = TRUE, out.width = '520'}

hex_tidyverse <- readJPEG('images/tidyverse.jpg', native = TRUE)
hex_tidyverse <- rasterGrob(hex_tidyverse, interpolate = FALSE)

hex_readr <- readJPEG('images/readr.jpg', native = TRUE)
hex_readr <- rasterGrob(hex_readr, interpolate = FALSE)

hex_tidyr <- readJPEG('images/tidyr.jpg', native = TRUE)
hex_tidyr <- rasterGrob(hex_tidyr, interpolate = FALSE)

hex_dplyr <- readJPEG('images/dplyr.jpg', native = TRUE)
hex_dplyr <- rasterGrob(hex_dplyr, interpolate = FALSE)

hex_ggplot2 <- readJPEG('images/ggplot2.jpg', native = TRUE)
hex_ggplot2 <- rasterGrob(hex_ggplot2, interpolate = FALSE)

hex_rmarkdown <- readJPEG('images/rmarkdown.jpg', native = TRUE)
hex_rmarkdown <- rasterGrob(hex_rmarkdown, interpolate = FALSE)

data.frame(
  x = c(2:5.0), y = c(1.6:4.6)
) %>%
  ggplot(aes(x, y)) +
  geom_blank() +
  annotate(geom = 'rect', xmin = 2.95, xmax = 3.93, ymin = 2.5, ymax = 3.6, colour = 'yellow', fill = 'yellow', alpha = 0.4) +
  annotate(geom = 'text', x = 2.2, y = 3.0, label = 'Importar', family = 'gochi', size = 10.4, colour = 'black') +
  geom_curve(x = 2.4, xend = 2.5, y = 3.0, yend = 3.0, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = 0.0) +
  annotate(geom = 'text', x = 2.7, y = 3.0, label = 'Ordenar', family = 'gochi', size = 10.4, colour = 'black') +
  geom_curve(x = 2.9, xend = 3.0, y = 3.0, yend = 3.0, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = 0.0) +
  annotate(geom = 'text', x = 3.25, y = 3.0, label = 'Transformar', family = 'gochi', size = 10.4, colour = 'black') +
  annotate(geom = 'text', x = 3.7, y = 3.5, label = 'Visualizar', family = 'gochi', size = 10.4, colour = 'black') +
  geom_curve(x = 3.4, xend = 3.6, y = 3.1, yend = 3.4, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = -0.2) +
  annotate(geom = 'text', x = 3.75, y = 2.6, label = 'Modelar', family = 'gochi', size = 10.4, colour = 'black') +
  geom_curve(x = 3.65, xend = 3.4, y = 2.7, yend = 2.9, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = -0.2) +
  geom_curve(x = 3.8, xend = 3.8, y = 3.4, yend = 2.7, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = -0.3) +
  geom_curve(x = 3.96, xend = 4.06, y = 3.0, yend = 3.0, arrow = arrow(length = unit(0.1, 'inch')), size = 1.0, color = 'black', curvature = 0.0) +
  annotate(geom = 'text', x = 4.3, y = 3.0, label = 'Comunicar', family = 'gochi', size = 10.4, colour = 'black') +
  annotate(geom = 'text', x = 2.98, y = 2.4, label = 'Explorar', family = 'gochi', size = 10.4, colour = 'black') +
  annotation_custom(hex_tidyverse, xmin = 4.0, xmax = 5.0, ymin = 3.8, ymax = 4.8) +
  annotation_custom(hex_readr, xmin = 2.1, xmax = 2.3, ymin = 3.1, ymax = 3.5) +
  annotation_custom(hex_tidyr, xmin = 3.1, xmax = 2.3, ymin = 3.1, ymax = 3.5) +
  annotation_custom(hex_dplyr, xmin = 4.1, xmax = 2.3, ymin = 3.14, ymax = 3.54) +
  annotation_custom(hex_ggplot2, xmin = 3.5, xmax = 3.9, ymin = 3.64, ymax = 4.04) +
  annotation_custom(hex_rmarkdown, xmin = 4.1, xmax = 4.5, ymin = 2.5, ymax = 2.9) +
  scale_x_continuous(limits = c(2, 4.6)) +
  theme_void() +
  theme(axis.text = element_blank(),
        axis.ticks = element_blank(),
        axis.title = element_blank())
```

A lo largo del curso podrá encontrar ejemplos de código y ejercicios de codificación. Abajo encontrara como podrán ser usados:

#### `r fa("fas fa-chalkboard-teacher", fill = "black")` = Código de ejemplo.

Los ejemplos contienen código previamente escrito para que pueda explorarlo y ejecutarlo presionando el botón amarillo "Run Code" (ejecutar) en la esquina superior derecha. Puede actualizar el código en cualquiera de los ejemplos presionando "Start Over" (iniciar de nuevo) en la esquina superior izquierda y el código se restablecerá automáticamente al valor predeterminado.

```{r addition_ejemplos, exercise = TRUE}


```

#### `r fa("fas fa-laptop-code", fill = "black")` = Ejercicio de codificación.

Aquí se indican ejercicios prácticos de codificación donde se le pedirá que escriba el código manualmente. Si no da con la solución, puede buscar ayuda dando clic en los botones "Solution" (solución) o "Hints" (sugerencias). En la mayoría de los casos el código podrá ser evaluado con el botón “Submit Answer” (Enviar Respuesta), indicándole cuando la respuesta sea acertada.

```{r exercise_ejercicios, exercise = TRUE}

```

```{r exercise_ejercicios-solution}

Ejercicio_código
```

<br>

### Acerca de nosotros

<div class="row">
  <div class="col-md-6" markdown="1">
  __Jorge Leonardo López Martínez__ es zootecnista con interés en el uso de la programación aplicada al campo de la genética animal.
  <HR>
  <i class="fas fa-envelope"></i> jollopezma@unal.edu.co<br>
  <i class="fa fa-github"></i> [Leo4Luffy](https://github.com/Leo4Luffy)
  </div>
  <div class="col-md-6" markdown="1">
  __Duvan Ariel Nieves Ruiz__ es estudiante de ingeniería ambiental apasionado por el software libre con interés en la ecología y la programación.
  <HR>
  <i class="fas fa-envelope"></i> dnieves@unal.edu.co<br>
  <i class="fa fa-github"></i> [Duvancho321](https://github.com/Duvancho321)
  </div>
</div>

<br>

No dude en comunicarse con nosotros por correo electrónico si tiene preguntas acerca del curso y su contenido.

Tenga en cuenta que este curso está bajo la licencia <a rel="license" href="http://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License</a>. Por tanto el contenido del curso debe acreditarse a sus autores.

<section style="text-align: center;"><a rel="license" href="http://creativecommons.org/licenses/by-nc-sa/4.0/"><img alt="Creative Commons License" style="border-width:0" src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" /></a><br/></section>

## 1. Una breve introducción al R

### ¿Qué es el `R`?

`R` es un lenguaje de programación creado por __Robert Gentleman__ y __Ross Ihaka__ en el año de 1992. Ambos creadores le dieron el nombre de `R` al lenguaje implementado por las iniciales de sus nombres (a modo de broma).

```{r, echo = FALSE, eval = TRUE, out.width = "28%", fig.align = "center"}

knitr::include_graphics("images/Creadores.jpg")
```

Se trata de un lenguaje __interpretado__ o de script, con __tipado dinámico__, __multiplataforma__ y __orientado a objetos__:

+ __Lenguaje interpretado__ o de script: se ejecuta utilizando
un programa intermedio llamado __intérprete__ (IDE).

+ __Orientado a objetos__: `R` realiza acciones sobre esos objetos (un número, un conjunto de datos, una palabra o un resumen estadístico como una media o varianza).

#### `r fa("fas fa-chalkboard-teacher", fill = "black")` Ejemplo 1.1

```{r addition_3, exercise = TRUE}

# Para asignar un tipo de dato a un objeto se suele emplear el simbolo "<-"

a <- 12

b <- c(4, 6, 8, 12)

c <- matrix(data = 1:4,
            nrow = 2,
            ncol = 2)
```

+ __Multiplataforma__: el intérprete de `R` está disponible en muchas plataformas (Linux, Windows, UNIX, Mac OS, etc), por lo que podrá ejecutarse en muchos computadores.

+ __Tipado dinámico__: se refiere a que no es necesario
declarar el __tipo de dato__. A su vez, `R` es __fuertemente tipado__.

#### `r fa("fas fa-chalkboard-teacher", fill = "black")` Ejemplo 1.2

```{r addition, exercise = TRUE}

a <- 'Hola'
typeof(a)

b <- 1L
typeof(b)

c <- 12.3
typeof(c)
```

### ¿Por qué usar `R`?

El lenguaje de programación `R` se puede obtener y distribuir de forma gratuita, debido a que se encuentra bajo __Licencia Pública General__ del proyecto GNU. Por lo tanto es un programa de código abierto y gratis.

<div class="row">
  <div class="col-md-8" markdown="1">
  `R` cuenta con una __comunidad__ de usuarios alrededor del mundo. Al contar con una comunidad es posible dar con la solución de algún problema de programación, así como el desarrollo de paquetes utilizados en temas específicos, desarrolados por especialistas en alguna parte del mundo.
  </div>
  <div class="col-md-4" markdown="1">
```{r, echo = FALSE, eval = TRUE, out.width = "54%", fig.align = "center"}

knitr::include_graphics("images/Rday.jpg")
```
<section style="text-align: center;"><small>Encuentro en el [Rday 2019](https://rdaymedellin.github.io/).</small></section>
  </div>
</div>

Finalmente, el uso del `R` garantiza que otro investigador pueda repetir el experimento, comprobar los resultados obtenidos y estar en la condición de ampliar o refutar las interpretaciones del estudio realizado.

<!--es importante resaltar que la reproducibilidad es uno de los principios de la investigación científica. E-->

```{r, echo = FALSE, eval = TRUE, out.width = "30%", fig.align = "center"}

knitr::include_graphics("images/al_principio_pero_ahora.jpg")
```

<section style="text-align: center;"><small>Ilustracíon hecha por <font color = "black"><u>@allison_horst</u></font>.</small></section>

### Instalación de `R` base y `RStudio`

`R` base es el software básico que contiene el lenguaje de programación `R`. `RStudio` es un software que facilita la programación en `R`.

#### `R` base:

<div class="row">
  <div class="col-md-6" markdown="1">
  + Windows: la instalación de `R` base en windows se realiza a través de la [CRAN](https://cran.r-project.org/mirrors.html) (Comprehensive R Archive Network).
  + Linux: la instalación de `R` base en Linux se realiza escribiendo el siguiente código en la terminal:<br>
  <div style="background-color:#262626">
  <font color = "white"><p> sudo apt-get update</p></font>
  <font color = "white"><p> sudo apt-get install r-base r-base-dev</p></font>
  </div>
  </div>
  <div class="col-md-6" markdown="1">
```{r, echo = FALSE, eval = TRUE, out.width = "54%", fig.align = "center"}

knitr::include_graphics("images/consola_R_base.jpg")
```
  </div>
</div>

#### `RStudio`:

<div class="row">
  <div class="col-md-6" markdown="1">
  Se puede instalar directamente desde la [página de `RStudio`](https://rstudio.com/products/rstudio/download/). Hay se encuentran los archivos de instalación tanto en windows como en Linux.
  
  Cuando se abre `RStudio` se pueden ver cuatro paneles: 
  </div>
  <div class="col-md-6" markdown="1">
```{r, echo = FALSE, eval = TRUE, out.width = "34%", fig.align = "center"}

knitr::include_graphics("images/RStudio.jpg")
```
  <section style="text-align: center;"><small>Ilustracíon hecha por <font color = "black"><u>@allison_horst</u></font>.</small></section>
  </div>
</div>

+ La fuente: su libreta de notas para el código

+ La consola: el corazón de `R`

+ El ambiente / el historial

+ Los archivos / las gráficas / los paquetes / la ayuda

<!--colocar sobre rscloud-->

## 2. Operaciones básicas R

Antes de realizar operaciones sobre algunos objetos en R, veamos cuales de estas podemos realizar.

* Suma y diferencia (+, -)
* Producto y división (*, /)
* Potencia (^,**)
* Residuo (%%)
* Mayor, mayor o igual, menor, menor o igual, igual y diferente que (>, >=, <, <=, ==, !=)
* Conjunción (&, &&), disyunción (|, ||)

### Operaciones con vectores 

Generemos un vector de prueba y realicemos algunas operaciones aritméticas básicas.Sean estas suma, resta, división, multiplicación y exponenciación.

Operaciones de suma.

#### `r fa("fas fa-chalkboard-teacher", fill = "black")` Ejemplo 2.1

```{r addition_1, exercise = TRUE}
#Vector de prueba
x <- c(1:8)
x

#Suma sobre cada elemento del vector de prueba
test_1 <- x+2
test_1

#Suma sobre repetitiva sobre los elementos del vector de prueba
test_2 <- x + c(10,20)
test_2

#Suma uno a uno sobre nuestro vector de prueba 
test_3 <- x + c(8:1) 
test_3
```

De igual forma podemos utilizar la resta, te invito a que restablezcas el vector de prueba `x` a partir de los vectores test_1, test_2, test_3. Para ello copia los vectores que realizamos en el fragmento de código anterior y crear de nuevo el vector `x`. Vamos confio en que podrás hacerlo.

#### `r fa("fas fa-laptop-code", fill = "black")` Ejercicio 2.1

```{r exe_1, exercise = TRUE}
test_1 <- c(1:8)+ 2
test_2 <- c(1:8) + c(10,20)
test_3 <- c(1:8) + c(8:1)

# Ahora es tu turno
```

```{r exe_1-solution}
test_1 - 2
test_2 - c(10,20)
test_3 - c(8:1)
```

```{r exe_1-check}
grade_code("Excelente lo lograste, sigamos adelante")
```

Vamos ahora con operación de producto. 

#### `r fa("fas fa-chalkboard-teacher", fill = "black")` Ejemplo 2.2

```{r addition_2, exercise = TRUE}
#Vector de prueba
y <- seq(2,20,2)
y

#Producto sobre cada elemento del vector de prueba
test_1 <- y*2
test_1

#Producto repetitivo sobre los dos elementos del vector de prueba
test_2 <- y * c(5,7)
test_2

#Producto uno a uno sobre nuestro vector de prueba 
test_3 <- y * c(10:1) 
test_3
```

Ya que vimos cómo funciona el producto en R, te propongo que utilizando la operación inversa restablezcas el vector de prueba `y`. Intentalo, se que podrás hacerlo.

#### `r fa("fas fa-laptop-code", fill = "black")` Ejercicio 2.2

```{r exe_2, exercise = TRUE}
test_1 <- seq(2,20,2) * 2
test_2 <- seq(2,20,2) * c(5,7)
test_3 <- seq(2,20,2) * c(10:1) 

# Ahora es tu turno
```

```{r exe_2-solution}
test_1 / 2
test_2 / c(10,20)
test_3 / c(10:1)
```

```{r exe_2-check}
grade_code("Excelente lo lograste, sigamos adelante")
```

Hagamos un ejercicio un poco más corto para la potenciación y su operación inversa. Halla la séptima potencia del 4 y la raíz cúbica del 27, te dejamos un par de ejemplos para que puedas guiarte. 

#### `r fa("fas fa-laptop-code", fill = "black")` Ejercicio 2.3

```{r exe_3, exercise = TRUE}
8**5
16**(1/4)

# Ahora es tu turno
```

```{r exe_3-solution}
4**7
27**(1/3)
```

```{r exe_3-check}
grade_code("Excelente lo lograste, sigamos adelante")
```

Ahora hagamos un par de ejercicios de operaciones lógicas, los cuales son muy útiles  en cualquier campo de la programación.


